# Review module

Throughout this foundational pre-program, you will complete several auto-graded review modules designed to test your understanding of the previous cases.

**Instructions**

The exercises will require you to write simple snippets of code. In order to complete this part, we recommend you follow these instructions:

1. Complete the code provided to you in this notebook, but do **not** change the name of the variables. If you do that, the autograder will fail and you will not receive any points.
2. To check that your answers are correct, click on `Kernel > Restart & Run All`. If any of the cells in your notebook produces an error, that means you need to adjust your answers.

In the questions below, you'll see a couple of lines that look like this:

~~~python
# YOUR CODE HERE
raise NotImplementedError() # Remove this line when you enter your solution
~~~

These lines are there to help us keep track of which exercises have been answered and which haven't. For the exercises in this review module, you should simply remove them and proceed to answering the exercises just as we explained above.

Here are a couple of already-solved example questions:

### Example 1

Write code that takes a numerical value, multiplies it by itself, and then outputs the result as the `result` variable.

```python
input_integer = 6
result = input_integer * input_integer
```

You can check that your variable was assigned by calling it like this:

```python
result
```

### Example 2

Write code that takes any integer value, adds it to its own value divided by 10, and outputs the result as the `result` variable.

```python
input_number = 30
result = input_number + (input_number/10)
```

And checking:

```python
result
```

Alright? Here come the coding questions!

### Exercise 1

Write code that finds the product of these 4 variables (do not change them!). Assign the result to a **float** variable called `ex_1_result`.

<!-- grade_id: cell-6da7fa29a659afd3, points: 0, grade: false, solution: true -->
```python
a = 5
b = 6
c = 7
d = 8
# Call the result variable ex_1_result
# YOUR CODE HERE
ex_1_result = float(a * b * c * d)
print(ex_1_result)
```

Output:
```
1680.0
```

### Exercise 2

Write code that adds (concatenates) these three strings together. Do **not** change the values of the variables. Call the result `ex_2_result`.

<!-- grade_id: cell-b74d0ff90247b40b, points: 0, grade: false, solution: true -->
```python
string1 = "Data"
string2 = "Skills for"
string3 = " All"
# Call the result variable ex_2_result
# YOUR CODE HERE
ex_2_result = string1 + string2 + string3
print(ex_2_result)
```

Output:
```
DataSkills for All
```

### Exercise 3

Write code that

* takes the string variable `my_string`,
* converts it to an integer,
* then divides that integer by the length of the original string,
* and outputs the result.

For example, if the `my_string` variable were `"567"`, then the result should be the number `189`, because `567/3=189`. Make sure that the result is of type **integer** and is assigned to a variable called `ex_3_result`. 

<!-- grade_id: cell-466de4b7eb3a9f70, points: 0, grade: false, solution: true -->
```python
my_string = "36189912"
# Call the result variable ex_3_result
# YOUR CODE HERE
my_string_integer = int(my_string)
ex_3_result = int(my_string_integer / len(my_string))
print(ex_3_result)
```

Output:
```
4523739
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-e3ff00e471bec774, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert "ex_1_result" in globals(), "Ex. 1 - You haven't defined the ex_1_result variable!"
assert type(ex_1_result) == type(0.5), "Ex. 1 - The ex_1_result variable is not a float! You can convert it with the float() function."
print("Preliminary sanity check for Ex. 1 passed.")
assert ex_1_result == 1680.0, "Ex. 1 - The result is wrong. It should be 1680.0 and it is " + str(ex_1_result) + " instead."
print("Exercise 1 looks correct!")
```

Output:
```
Preliminary sanity check for Ex. 1 passed.
Exercise 1 looks correct!
```

<!-- grade_id: cell-c0480287df6d8dd5, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert "ex_2_result" in globals(), "Ex. 2 - You haven't defined the ex_2_result variable!"
print("Preliminary sanity check for Ex. 2 passed.")
assert ex_2_result == 'DataSkills for All', "Ex. 2 - The result is wrong. It should be 'DataSkills for All' and it is '" + ex_2_result + "' instead."
print("Exercise 2 looks correct!")
```

Output:
```
Preliminary sanity check for Ex. 2 passed.
Exercise 2 looks correct!
```

<!-- grade_id: cell-bbb0af688a1c570a, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert "ex_3_result" in globals(), "Ex. 3 - You haven't defined the ex_3_result variable!"
assert type(ex_3_result) == type(5), "Ex. 3 - The ex_3_result variable is not an integer! You can convert it with the int() function."
print("Preliminary sanity check for Ex. 3 passed.")
assert ex_3_result == 4523739, "Ex. 3 - The result is wrong. It should be 4523739 and it is " + str(ex_3_result) + " instead."
print("Exercise 3 looks correct!")
```

Output:
```
Preliminary sanity check for Ex. 3 passed.
Exercise 3 looks correct!
```