# Review module

Throughout this foundational pre-program, you will complete several auto-graded review modules designed to test your understanding of the previous cases.

**Instructions**

For this review module, the exercises will require you to uncomment lines of code. In order to complete this part, we recommend you follow these instructions:

1. Complete the code provided to you in this notebook, but do **not** change the name of the variables. If you do that, the autograder will fail and you will not receive any points.
2. To check that your answers are correct, click on `Kernel > Restart & Run All`. If any of the cells in your notebook produces an error, that means you need to adjust your answers.

Here are a couple of already-solved example questions:

### Example 1

What is the coolest animal in the world? (Remember, you need to uncomment the line that corresponds to the correct answer - you uncomment by removing the `#` symbol).

* A. Wombat
* B. Red panda.
* C. Hyrax.
* D. Meerkat.

```python
#answer = "A"
#answer = "B"
#answer = "C"
answer = "D"
```

### Example 2

Why did the chicken cross the road? (Remember, you need to uncomment the line that corresponds to the correct answer - you uncomment by removing the `#` symbol).

* A. To attend an astrophysics conference.
* B. To chat with a cool meerkat.
* C. To get to the other side.
* D. To have its hair done.

```python
#answer = "A"
answer = "B"
#answer = "C"
#answer = "D"
```

In the questions below, you'll see a couple of lines that look like this:

~~~python
# YOUR CODE HERE
raise NotImplementedError() # Remove this line when you enter your solution
~~~

These lines are there to help us keep track of which exercises have been answered and which haven't. For the exercises in this review module, you should simply remove them and proceed to uncommenting the lines just as we explained above.

Alright? Here come the questions!

### Exercise 1

In computer science, Python is a(n)...

* A. IDE
* B. function
* C. programming language
* D. type of snake

<!-- grade_id: cell-19572b3e477e143f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

#ex_1_answer = "A"
#ex_1_answer = "B"
ex_1_answer = "C"
#ex_1_answer = "D"
```

### Exercise 2

The 3 disciplines necessary to become a great data professional are...

* A. Machine Learning, AI, Data Science

* B. Math & Statistics, Computer Science, Business Expertise

* C. Math & Statistics, Computer Programming, Data Science

* D. Math & Statistics, Python, Business Expertise

<!-- grade_id: cell-9f31b59057e28bc8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

#ex_2_answer = "A"
ex_2_answer = "B"
#ex_2_answer = "C"
#ex_2_answer = "D"
```

### Exercise 3

The `cd` command lets you...

* A. Rename folders
* B. See the contents of a folder
* C. Create a new folder
* D. Navigate around your computer's folders

<!-- grade_id: cell-b34688694bf03284, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

#ex_3_answer = "A"
#ex_3_answer = "B"
#ex_3_answer = "C"
ex_3_answer = "D"
```

### Exercise 4

The `ls` command allows you to...

* A. see the name of the current file
* B. look into the contents of a file
* C. see the contents of a folder
* D. list the size of the current file

<!-- grade_id: cell-a7ef3f10e0b79b90, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

#ex_4_answer = "A"
#ex_4_answer = "B"
ex_4_answer = "C"
#ex_4_answer = "D"
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-dfce3810f2c638eb, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert "ex_1_answer" in globals(), "Ex. 1 - You haven't defined the ex_1_answer variable!"
print("Preliminary sanity check for Ex. 1 passed.")
assert ord(ex_1_answer) == 67, "Ex. 1 - Think again! Check the main case for a refresher!"
print("Exercise 1 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 1 passed.
Exercise 1 seems correct!
```

<!-- grade_id: cell-f0e6846c9b920e1e, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert "ex_2_answer" in globals(), "Ex. 2 - You haven't defined the ex_1_answer variable!"
print("Preliminary sanity check for Ex. 2 passed.")
assert ord(ex_2_answer) == 66, "Ex. 2 - Think again! Check the main case for a refresher!"
print("Exercise 2 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 2 passed.
Exercise 2 seems correct!
```

<!-- grade_id: cell-71380078e8854049, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert "ex_3_answer" in globals(), "Ex. 3 - You haven't defined the ex_3_answer variable!"
print("Preliminary sanity check for Ex. 3 passed.")
assert ord(ex_3_answer) == 68, "Ex. 3 - Think again! Check the main case for a refresher!"
print("Exercise 3 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 3 passed.
Exercise 3 seems correct!
```

<!-- grade_id: cell-1c9d233c0f481434, points: 1, grade: true, solution: false -->
```python
# Ex. 4
assert "ex_4_answer" in globals(), "Ex. 4 - You haven't defined the ex_4_answer variable!"
print("Preliminary sanity check for Ex. 4 passed.")
assert ord(ex_4_answer) == 67, "Ex. 4 - Think again! Check the main case for a refresher!"
print("Exercise 4 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 4 passed.
Exercise 4 seems correct!
```